```python
# Initialize Otter
import otter
grader = otter.Notebook("Project.ipynb")
```

# Jupyter Notebook and Autograder Tutorial

The goal of this project is to refresh your memory on how to use Jupiter notebooks, did you familiar with the auto greeting process will be using in the course, and to iron out any technical issues maybe having.

<!-- BEGIN QUESTION -->

Write a function that returns an empty Pandas DataFrame

```python
import pandas as pd
def make_dataframe():
    ...
```

```python
grader.check("q1")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Write a function that uses Python's [`sqlite3` module](https://docs.python.org/3/library/sqlite3.html) to create a table called `college_mascots`. Your table must contain at least one field called `mascot_name` but can contain any other fields you would like. The function will receive an [`Connection`](https://docs.python.org/3/library/sqlite3.html#sqlite3.Connection) object as an argument; your table must be created using that object.

```python
import sqlite3
def make_college_mascots_table(conn: sqlite3.Connection):
    ...
```

```python
grader.check("q2")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

We're going to connect to MongoDB next! But first, we need to start MongoDB

### Starting MongoDB

Open a terminal in Jupyter Lab by clicking the big "plus" button

<img src="img/make-terminal-1.png" alt="click blue button" width="200"/>

then click the "Terminal" button to launch a terminal.

<img src="img/make-terminal-2.png" alt="click Terminal button" width="200"/>

Finally, type `start-mongo` and hit Enter before returning to this notebook

### Your Task

Write a function which uses the `pymongo` Python client library for MongoDB to create a new *database* called `music_library` containing a new *collection* called `albums`. The function should return the `albums` collection.

Refer to the [`pymongo` Documentation](https://pymongo.readthedocs.io/en/stable/index.html) when necessary

```python
def make_album_collection():
    import pymongo
    client = pymongo.MongoClient()
    ...
```

```python

```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

These are some submission instructions.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> assert make_dataframe() is not None, "Did not return anything"
>>> assert isinstance(make_dataframe(), pd.DataFrame), "Did not return a DataFrame instance"
```


Tests for q2 (points: None)
```
>>> conn = sqlite3.connect(':memory:') #  IGNORE
>>> make_college_mascots_table(conn)
>>> res = conn.execute("select 1 from sqlite_master where type='table' and name='college_mascots';").fetchone()
>>> assert res is not None and len(res) == 1, "No table named 'college_mascots'"
>>> res = conn.execute("SELECT 1 FROM PRAGMA_TABLE_INFO('college_mascots') where name='mascot_name';").fetchone()
>>> assert res is not None and len(res) == 1, "No field named 'mascot_name' in 'college_mascots' table"
```


Tests for q3 (points: None)
```
>>> col = make_album_collection()
>>> assert col.database.name == 'music_library'
>>> assert col.name == 'albums'
```
